# Problem Set 1

## Question 1: R Tutorial

Work through the provided tutorial (FirstRTutorial.pdf). Type in and run all the example code as you read through the document. This is especially important for first-time coders, as it will help you learn the proper notation. You should complete all the exercises in the handout, but you only need to hand in the ones listed below. 
 
I. Section 3, Problem 2 

II. Section 5, Problems 1 and 2 

III. Section 7, Problem 1 

IV. Section 8, Problem 1 

V. Section 9, Problems 1 and 2 

A hint for III: remember that $X \ \text{~} \ \mathcal{N}(3,2^2)$ just means that $X$ is normally distributed with a mean of 3 and a standard deviation of 2. If you need to brush up on the terminology or mathematics of distributions please refer to the additional materials section of this week’s assignments. For more details on these R functions, see http://www.cyclismo.org/tutorial/R/probability.html.


```R
#attach to datasets -- only run this cell once!

data(quakes)
attach(quakes)

data(mtcars)
attach(mtcars)
```

<!-- grade_id: r_tutorial_Qs, points: 0, grade: false, solution: true -->
```R
# I

y <- c(33,44,29,16,25,45,33,19,54,22,21,49,11,24,56)

day2 <- mean(y[4:6]) #store summary stats for the second day here
shop3 <- mean(y[seq(3,length(y),by=3)]) #store summary stats for the third shop here
#NOTE: only these answers will be checked

# II

depthMean <- mean(depth)
depthMedian <- median(depth)
magMin <- min(mag)
magMax <- max(mag)

meanWeight <- mean(wt)
meanFuel <- mean(mpg)

# III

a <- dnorm(0.5,2,0.5)
b <- pnorm(2.5,2,0.5)
c <- qnorm(0.95,2,0.5)
d <- pnorm(3,2,0.5)-pnorm(1,2,0.5)
```

<!-- grade_id: r_tutorial_Q4, points: 1, grade: true, solution: true -->
```R
# IV
```

<!-- grade_id: r_tutorial_Q5, points: 0, grade: false, solution: true -->
```R
# V

myCorPlotFunk <- function(x,y){
    return cor(x,y)
} #make sure to only return the correlation! 
#if you're curious, check what happens if you return the output of a plotting function

mySumsFunk <- function(x) {
    return c(sum(x),sum(x^2))
} #return the sum and the sum of squares in that order!
```

<!-- grade_id: r_tutorial_Q5_tests, points: 1, grade: true, solution: false -->
```R
x = rnorm(10); y = rnorm(10)
assert_that(are_equal(myCorPlotFunk(x,y),cor(x,y)))
assert_that(are_equal(mySumsFunk(x),c(sum(x),sum(x^2))))
```

## Question 2: What's in a norm?

Explain the `dnorm`, `pnorm`, and `qnorm` functions in your own words.

<!-- grade_id: question_2, points: 3, grade: true, solution: true -->
```R

```

## Question 3: Time Flies

You’re studying courtship behavior in flies and you want to know the average number of times a wildtype male fly attempts copulation in one day. You’ve videotaped 30 individual male flies in vials with dozens of virgin females. After you’ve subjected your undergrad to watching all the videos and counting the copulation attempts, she provides you with the numbers below.

First, calculate the mean, median, and standard deviation of your sample. Produce a histogram with six bins and six axis labels.
Then, create a function that computes a 95% confidence interval from an input vector and calculate a 95% confidence interval on the data. 
Finally, explain, in your own words, what this means. Answer in the text box below.

<!-- grade_id: question_3, points: 1, grade: true, solution: true -->
```R
flyData<-c(31,15,12,21,26,10,12,25,27,25,
16,19,16,18,8,13,22,19,29,17,
30,14,9,22,28,19,30,23,20,17)

flyMean <- mean(flyData)
flyMedian <- median(flyData)
flySD <- sqrt(var(flyData))

err <- qnorm(0.975)*(flySD)/sqrt(length(flyData))

confInterval <- c(flyMean-err,flyMean+err)

myConfFunk <- function(x) {} #return confInterval as a vector, lower bound then upper

#ALSO PRINT YOUR HISTOGRAM IN THIS BOX
```

```R

```

<!-- grade_id: conf_interval, points: 2, grade: true, solution: true -->
A 95% confidence interval is